##### Test question (just for you to check your environment is ready and familiar with writing code inside the function)

Write a function that returns a list of numbers, such that  ğ‘¥ğ‘–=ğ‘–2 , for  1â‰¤ğ‘–â‰¤ğ‘› . Make sure it handles the case where  ğ‘›<1  by raising a ValueError.

<!-- grade_id: cell-b4b8af58225a3d15, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1c8f4d2067ce3c82, points: 0, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

```python
###answer###
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    if n < 1:
        raise ValueError("n must be greater than or equal to 1")
    return [i ** 2 for i in range(1, n + 1)]
```

# 1.

<!-- grade_id: cell-cb03e01b962b066b, points: 0, grade: false, solution: true -->
```python
def Q1():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3ff2a4b067a8451a, points: 4, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a5d24c3f856a7d82, points: 4, grade: true, solution: false -->
```python

```

# 2.

<!-- grade_id: cell-8f3bf76a6f5b7386, points: 0, grade: false, solution: true -->
```python
from random import randint
rn=[randint(1,1500) for _ in range (50)]
def Q2(rn):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a21662173603ccd0, points: 8, grade: true, solution: false -->
```python

```

# 3.

<!-- grade_id: cell-1e67a55cc8f3479b, points: 0, grade: false, solution: true -->
```python
def Q3():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4b1c46ec4b2dcf2b, points: 8, grade: true, solution: false -->
```python

```

# 4.

<!-- grade_id: cell-f0a54259e3a0bf46, points: 0, grade: false, solution: true -->
```python
def Q4():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f89e6dc23ece3507, points: 4, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-227dd601a021567a, points: 4, grade: true, solution: false -->
```python

```

# 5.

<!-- grade_id: cell-62813295633b87e9, points: 0, grade: false, solution: true -->
```python
def Q5():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-77c0d621ffa15f74, points: 4, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-7723c1278cd2ed97, points: 4, grade: true, solution: false -->
```python

```

# 6.

<!-- grade_id: cell-b6adb8cddc912730, points: 0, grade: false, solution: true -->
```python
def Q6():
    from random import randint
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d4c2cee57c3020d8, points: 8, grade: true, solution: false -->
```python

```

# 7.

<!-- grade_id: cell-928b75b0295ad37a, points: 0, grade: false, solution: true -->
```python
def Q7():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b288faf2b4eac1fb, points: 10, grade: true, solution: false -->
```python
#
```

# 8.

<!-- grade_id: cell-39c21130c38e18eb, points: 0, grade: false, solution: true -->
```python
def Q8(title):
    # YOUR CODE HERE
    raise NotImplementedError()
print(f'Capitalized: {Q8(input("Enter a title: "))}')
```

<!-- grade_id: cell-0f4a0cbb0b5df331, points: 4, grade: true, solution: false -->
```python
assert Q8('hi') == 'Hi'
```

<!-- grade_id: cell-56086f41becc18ad, points: 3, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-41b3501d029296dd, points: 10, grade: true, solution: false -->
```python

```

# 9.

<!-- grade_id: cell-30dd91b3d407935e, points: 0, grade: false, solution: true -->
```python
import math, pandas
def Q9():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-028c2007345ea430, points: 3, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a5f4630089bf96d4, points: 3, grade: true, solution: false -->
```python
#
```

<!-- grade_id: cell-1480afa2e8dba94d, points: 3, grade: true, solution: false -->
```python
#
```

# 10.

<!-- grade_id: cell-4945423fc6345dff, points: 0, grade: false, solution: true -->
```python
def Q10():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5f6d766e856b85c1, points: 4, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-b22bbca9b3c13750, points: 4, grade: true, solution: false -->
```python

```

# 11.

<!-- grade_id: cell-8bd88b5236bf7b1b, points: 0, grade: false, solution: true -->
```python
def Q11():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-77eb254b3e6549ee, points: 4, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-62883fc7476c669b, points: 4, grade: true, solution: false -->
```python

```

# 12.

<!-- grade_id: cell-55ff778b67c81a57, points: 0, grade: false, solution: true -->
```python
def Q12():
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fde7e2ebe96e0a4b, points: 0, grade: true, solution: false -->
```python

```

# 13.

<!-- grade_id: cell-a62dd681381d5eaa, points: 0, grade: false, solution: true -->
```python
def Q13():
     
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-26e496b93a127cfb, points: 0, grade: true, solution: false -->
```python

```

```python

```